# CS6140 Assignments

**Instructions**
1. In each assignment cell, look for the block:
 ```
  #BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
 ```
1. Replace this block with your solution.
1. Test your solution by running the cells following your block (indicated by ##TEST##)
1. Click the "Validate" button above to validate the work.

**Notes**
* You may add other cells and functions as needed
* Keep all code in the same notebook
* In order to receive credit, code must "Validate" on the JupyterHub server

---

```ruby

```

# Assignment 7: Kernels


```ruby
require './assignment_lib'
```

Output:
```
false
```

## Question 1.1: Prepare Dataset (1 points)

We will generate and plot a synthetic dataset to illustrate Kernel learning. Here is a pre-defined dataset for us to use in this exercise. Paste **your** dot and norm function in the cell below.

```ruby
def plot_spiral_dataset()
  spiral = spiral_dataset()["data"]
  x1 = spiral.map {|r| r["features"]["x1"]}
  x2 = spiral.map {|r| r["features"]["x2"]}
  target = spiral.map {|r| r["label"]}
  target.size  
  df = Daru::DataFrame.new({x1: x1, x2: x2, target: target})
  df.to_category :target
  df.plot(type: :scatter, x: :x1, y: :x2, categorized: {by: :target, method: :color}) do |plot, diagram|
    plot.xrange [-8,8]
    plot.x_label "X1"
    plot.yrange [-8,8]  
    plot.y_label "X2"
    plot.legend false
  end
end

plot_spiral_dataset()
```

Output:
```
#<Nyaplot::Plot:0x0000000002310c48 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000000003c79d10 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1.0, :color=>"rgb(228,26,28)", :tooltip_contents=>[1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0]}, :data=>"32fd3177-c5f0-4dee-aa10-73e00406d8b5"}, @xrange=[-5.0, 6.0], @yrange=[-5.5, 6.5]>, #<Nyaplot::Diagram:0x0000000003bba708 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0.0, :color=>"rgb(55,126,184)", :tooltip_contents=>[0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0,
```

<!-- grade_id: cell-ff4b2c6b485a3e03, points: 0, grade: false, solution: true -->
```ruby
# BEGIN YOUR CODE
#Implement the error function given a weight vector, w
def dot x, w
  # BEGIN YOUR CODE
  sum = 0.0
    
    if !(x.empty? or w.empty?)
      x.each do |k, v|
          if w.has_key?(k)
              sum += v * w[k]
          end
      end
    end
    
    return sum
  #END YOUR CODE  
end

def norm w
  # BEGIN YOUR CODE
  return Math.sqrt(dot(w, w))
  #END YOUR CODE
end
#END YOUR CODE
```

Output:
```
:norm
```

<!-- grade_id: cell-aabc301bebd79fdd, points: 1, grade: true, solution: false -->
```ruby
def test_11()
  assert_in_delta 2.0, norm({"a" => 1.41421, "b" => 1.41421}), 1e-2
  assert_in_delta 2.0, norm({"a" => -1.41421, "b" => 1.41421}), 1e-2
  assert_in_delta 0.0, norm({}), 1e-2

  assert_in_delta 6.0, dot({"a" => 2.0}, {"a" => 3.0}), 1e-6
  assert_in_delta 6.0, dot({"a" => 2.0}, {"a" => 3.0, "b" => 4.0}), 1e-6
  assert_equal 0.0, dot({}, {})
  assert_equal 0.0, dot({"a" => 1.0}, {"b" => 1.0})
end

test_11()
```

## Question 1.2 (1 Point)
Copy **YOUR** implementation of ```StochasticGradientDescent``` from [Assignment 5](../assignment-5/assignment-5.ipynb) into the following cell.

<!-- grade_id: cell-512b41c089c45d6b, points: 0, grade: false, solution: true -->
```ruby
# BEGIN YOUR CODE
class StochasticGradientDescent
  attr_reader :weights
  attr_reader :objective
  def initialize obj, w_0, lr = 0.01
    @objective = obj
    @weights = w_0
    @n = 1.0
    @lr = lr
  end
  def update x
    # BEGIN YOUR CODE
    g = @objective.grad(x, @weights)
    learning_rate = @lr / Math.sqrt(@n)
    @weights.each do |k, v|
      @weights[k] -= g[k] * learning_rate
    end
    @n += 1.0
    #END YOUR CODE
  end
end
#END YOUR CODE
```

Output:
```
:update
```

<!-- grade_id: cell-decbf7a78199df43, points: 1, grade: true, solution: false -->
```ruby
### Hidden Test (See Test 1.1 from Assignment 5) ###
assert_not_nil(StochasticGradientDescent.class)
```

## Question 1.3 (5 Points)


Refactor **your** implementation of the AUC metric to return the false positive, true positive values for all scores in addition to the AUC. We will use these to plot the ROC curve.

```ruby
def cal_area (fp_rate, tp_rate)
  area = 0.0
  
  x_prev = 0.0
  y_prev = 0.0
  fp_rate.length.times do |i|
    x = fp_rate[i]
    y = tp_rate[i]
    area += (y + y_prev) * (x - x_prev) / 2.0
    x_prev = x
    y_prev = y
  end
  return area
end
```

Output:
```
:cal_area
```

<!-- grade_id: cell-b97ece98e0d680af, points: 0, grade: false, solution: true -->
```ruby
def roc_curve(scores)
  # BEGIN YOUR CODE
  sorted_scores = scores.sort { |a,b| a[0] <=> b[0] }
  sorted_label = sorted_scores.collect { |s| s[1] }
  sorted_y_hat = sorted_scores.collect { |s| s[0] }
  positive = sorted_label.count { |l| l == 1 }.to_f
  negative = sorted_label.count { |l| l < 1 }.to_f
  
  res = []
  res << [0, 0, 0]
  tp = 0
  fp = 0
  (sorted_label.length - 1).downto(0) do |i|
    threshold = sorted_y_hat[i]
    if sorted_label[i] > 0
      tp += 1
    else
      fp += 1
    end
    fp_rate = negative == 0.0 ? 0.0 : fp / negative;
    tp_rate = positive == 0.0 ? 0.0 : tp / positive;
    res << [fp_rate, tp_rate, threshold]
  end
  curve_x = res.collect{ |r| r[0] }
  curve_y = res.collect{ |r| r[1] }
  auc = cal_area(curve_x, curve_y)
  #END YOUR CODE
  return [res[0], res[1], auc]
end

```

Output:
```
:roc_curve
```

```ruby
  good_model = [[0.9, 1], [0.89, 1], [0.7, 0], [0.8, 1], [0.8, 0], [0.7, 1], [0.6, 0], [0.5, 0], [0.1, 0]]
  assert_true(roc_curve(good_model).last > 0.8)
  assert_true(roc_curve(good_model).last < 1)
  
  srand(777)
  ok_model = Array.new(100) {|i| [100 - i, (rand < (100 - i) / 100.0) ? 1 : 0] }
  ok_auc = roc_curve(ok_model).last
  assert_in_delta(0.8631239935587761, ok_auc, 1e-3)
  
  bad_model = Array.new(1000) {|i| [1000 - i, rand < 0.5 ? 1 : 0] }
  bad_auc = roc_curve(bad_model).last
  assert_in_delta(0.5, bad_auc, 5e-2)
  
  fp, tp, auc = roc_curve(good_model)
```

Output:
```
[[0, 0, 0], [0.0, 0.25, 0.9], 0.9]
```

<!-- grade_id: cell-bce19441b35b0abd, points: 5, grade: true, solution: false -->
```ruby
def test_13()
  good_model = [[0.9, 1], [0.89, 1], [0.7, 0], [0.8, 1], [0.8, 0], [0.7, 1], [0.6, 0], [0.5, 0], [0.1, 0]]
  assert_true(roc_curve(good_model).last > 0.8)
  assert_true(roc_curve(good_model).last < 1)
  
  srand(777)
  ok_model = Array.new(100) {|i| [100 - i, (rand < (100 - i) / 100.0) ? 1 : 0] }
  ok_auc = roc_curve(ok_model).last
  assert_in_delta(0.8631239935587761, ok_auc, 1e-3)
  
  bad_model = Array.new(1000) {|i| [1000 - i, rand < 0.5 ? 1 : 0] }
  bad_auc = roc_curve(bad_model).last
  assert_in_delta(0.5, bad_auc, 5e-2)
  
  fp, tp, auc = roc_curve(good_model)
  
  plot = Daru::DataFrame.new({x: fp, y: tp}).plot(type: :line, x: :x, y: :y) do |plot, diagram|
    plot.x_label "False Positive Rate"
    plot.y_label "True Positive Rate"
    diagram.title("Good Model (AUC > 0.8)")
    plot.legend(true)
  end
  
  fp, tp, auc = roc_curve(bad_model)
  plot.add(:line, fp, tp).color("red").title("Bad Model (AUC = 0.5)")
  
  plot.show()
end

test_13()
```

Output:
```
#<Nyaplot::Frame:0x0000000003a9c830 @properties={:panes=>[#<Nyaplot::Plot:0x0000000003acc850 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000000003aca0c8 @properties={:type=>:line, :options=>{:x=>:x, :y=>:y, :title=>"Good Model (AUC > 0.8)"}, :data=>"cf9f0bfb-8d8b-485d-8a89-79dcb1f7dc2d"}, @xrange=[0, 0], @yrange=[0.0, 0.9]>, #<Nyaplot::Diagram:0x0000000003a9d0a0 @properties={:type=>:line, :options=>{:x=>"data0", :y=>"data1", :color=>"red", :title=>"Bad Model (AUC = 0.5)"}, :data=>"32439938-45f8-49ca-8c34-35c65e47acb4"}, @xrange=[0, 0], @yrange=[0.0, 1000]>], :options=>{:x_label=>"False Positive Rate", :y_label=>"True Positive Rate", :legend=>true, :zoom=>true, :width=>800, :xrange=>[0, 0], :yrange=>[0.0, 1000]}}>], :data=>{"cf9f0bfb-8d8b-485d-8a89-79dcb1f7dc2d"=>#<Nyaplot::DataFrame:0x0000000003aca4d8 @name="cf9f0bfb-8d8b-485d-8a89-79dcb1f7dc2d", @rows=[{:x=>0, :y=>0.0}, {:x=>0, :y=>0.25}, {:x=>0, :y=>0.9}]>, "32439938-45f8-49ca-8c34-35c65e47acb4"=>#<Nyaplot::DataFrame:0x0000000003a
```

##  Question 2.1 (10 Points)

Implement the Hinge Loss regularized linear model. The prediction function is the same as linear regression

<!-- grade_id: cell-2f4f3a643f0b678c, points: 0, grade: false, solution: true -->
```ruby
class HingeLossClassifier
  def initialize reg
    @reg_param = reg
  end
  
  def predict row, w
    # BEGIN YOUR CODE
    x = row["features"]    
    yhat = dot(w, x)
    #END YOUR CODE
  end
  
  def adjust w
  end
end
```

Output:
```
:adjust
```

<!-- grade_id: cell-67fe4fb894a07ef0, points: 10, grade: true, solution: false -->
```ruby
## Tests predict function
def test_21()
  t_loss = HingeLossClassifier.new 0.0
  assert_equal(1.0, t_loss.predict({"features" => {"x1" => 1.0}}, {"x1" => 1.0}))
  assert_equal(-1.0, t_loss.predict({"features" => {"x1" => -1.0}}, {"x1" => 1.0}))
  assert_equal(-14.0, t_loss.predict({"features" => {"x1" => 2.0}}, {"x1" => -7.0}))
end

test_21()
```

## Question 2.2 (10 Points)

Implement the objective function, which is defined as follows for a mini-batch of $n$ examples:

# $L(w,X) = \frac{\lambda}{2} \left\lVert w \right\rVert ^ 2 + \frac{1}{n} \sum_{i} \max \left[ 0, 1 - y_i \widehat{y}_i\right]$

where ```@reg_param``` is the value $\lambda$.

Remember that for hinge loss, the label must be $y = \pm 1$.

<!-- grade_id: cell-92a64a66671de794, points: 0, grade: false, solution: true -->
```ruby
class HingeLossClassifier
  def func data, w
    # BEGIN YOUR CODE
    l = 0.0
    data.each do |x|
      l += [0, 1 - predict(x, w) * (x["label"] > 0? 1 : -1)].max
    end
    obj = 0.5 * @reg_param * dot(w, w) + l / data.size
    #END YOUR CODE
    return obj
  end
end
```

Output:
```
:func
```

<!-- grade_id: cell-2dbf2caa8b69e62c, points: 10, grade: true, solution: false -->
```ruby
def test_22()
  ## Tests objective function
  t_loss = HingeLossClassifier.new 0.0
  t_data = [{"features" => {"x1" => 1.0}, "label" => 1.0}]
  assert_equal(0.0, t_loss.func(t_data, {"x1" => 1.0}), "1")
  assert_equal(2.0, t_loss.func(t_data, {"x1" => -1.0}), "2")
  assert_equal(1.0, t_loss.func(t_data, {"x1" => 0.0}), "3")

  t_data += [{"features" => {"x1" => 1.0}, "label" => 0.0}]
  assert_equal(1.0, t_loss.func(t_data, {"x1" => 1.0}), "4")
  assert_equal(1.0, t_loss.func(t_data, {"x1" => -1.0}), "5")
  assert_equal(1.0, t_loss.func(t_data, {"x1" => 0.0}), "6")

  t_loss = HingeLossClassifier.new 27.0
  t_data = [{"features" => {"x1" => 1.0}, "label" => 1.0}]
  assert_in_delta(675.0, t_loss.func(t_data, {"x1" => 1.0, "x2" => 7.0}), 1e-2, "7")
  assert_in_delta(677.0, t_loss.func(t_data, {"x1" => -1.0, "x2" => -7.0}), 1e-2, "8")
end
test_22()
```

## Question 2.3 (10 Points)

Implement the Hinge Loss classifier gradient method.


<!-- grade_id: cell-31495bf0a5d226c2, points: 0, grade: false, solution: true -->
```ruby
class HingeLossClassifier
  def grad data, w
    g = Hash.new {|h,k| h[k] = 0.0}
    # BEGIN YOUR CODE
    w.each_key do |key|
      g[key]
    end
    data.each do |x|
      y = x["label"] > 0 ? 1.0 : -1.0
      if y * predict(x, w) < 1
        x["features"].each_key do |f|
          g[f] -= (x["features"][f] * y)
        end
      end
    end
    g.each_key do |key|
      g[key] /= data.size
      g[key] += @reg_param * w[key]
    end
    #END YOUR CODE
    return g
  end
end
```

Output:
```
:grad
```

<!-- grade_id: cell-ce1477dcf71b4fbe, points: 10, grade: true, solution: false -->
```ruby
def test_33()
  t_loss = HingeLossClassifier.new 0.0
  t_data = [{"features" => {"x1" => 1.0}, "label" => 1.0}]
  assert_equal({"x1" => 0.0}, t_loss.grad(t_data, {"x1" => 1.0}), "1")
  assert_equal({"x1" => -1.0}, t_loss.grad(t_data, {"x1" => -1.0}), "2")
  assert_equal({"x1" => -1.0}, t_loss.grad(t_data, {"x1" => 0.0}), "3")

  t_data += [{"features" => {"x1" => 1.0}, "label" => 0.0}]
  assert_equal({"x1" => 0.5}, t_loss.grad(t_data, {"x1" => 1.0}), "4")
  assert_equal({"x1" => -0.5}, t_loss.grad(t_data, {"x1" => -1.0}), "5")
  assert_equal({"x1" => 0.0}, t_loss.grad(t_data, {"x1" => 0.0}), "6")

  t_loss = HingeLossClassifier.new 11.7
  t_data = [{"features" => {"x1" => 1.0}, "label" => 0.0}]
  assert_in_delta(12.7, t_loss.grad(t_data, {"x1" => 1.0})["x1"], 1e-2, "7")
  assert_in_delta(-11.7, t_loss.grad(t_data, {"x1" => -1.0})["x1"], 1e-2, "8")
  assert_in_delta(1.0, t_loss.grad(t_data, {"x1" => 0.0})["x1"], 1e-2, "9")

end
test_33()
```

## Question 2.4 (10 Points)

Train and test the Hinge Loss classifier on the same (training) sprial dataset. Return an array or (score, label) pairs. Try different learning rates, epochs, batch sizes, and regularization parameters. 

```ruby
def binary_classification(data, weights, model)
  score = Array.new (data.size) {Array.new(2, 0)}
  i = 0
  data.each do |x|
    pair = Array.new (2)
    y_hat = model.predict(x, weights)
    pair[0] = y_hat
    pair[1] = x["label"]
    score[i] = pair
    i += 1
  end
  return score
end
```

Output:
```
:binary_classification
```

<!-- grade_id: cell-e4ab95acd13266d2, points: 0, grade: false, solution: true -->
```ruby
def train_linear_hinge(dataset)
  lr = HingeLossClassifier.new 0.1
  w = Hash.new {|h,k| h[k] = (rand * 0.1) - 0.05}
  
  # BEGIN YOUR CODE
  sgd = StochasticGradientDescent.new lr, w, 0.3
  50.times do |t|
    sgd.update dataset["data"].sample(20)
  end
  scores = binary_classification(dataset["data"],w,lr)
  return scores 
  #END YOUR CODE
end
```

Output:
```
:train_linear_hinge
```

<!-- grade_id: cell-57a39fde4daf78e8, points: 10, grade: true, solution: false -->
```ruby
def test_24()
  spiral = spiral_dataset()
  scores = train_linear_hinge(spiral)
  fp, tp, auc = roc_curve(scores)
  puts auc

  assert_true(auc > 0.5, "AUC > 0.5")
  assert_true(auc > 0.52, "AUC > 0.52")  
  assert_true(auc < 0.7, "AUC < 0.7")    

  plot = Daru::DataFrame.new({x: fp, y: tp}).plot(type: :line, x: :x, y: :y) do |plot, diagram|
    plot.x_label "False Positive Rate"
    plot.y_label "True Positive Rate"
  end
  plot.add(:line, [0,1], [0,1]).color(:gray)
  plot.show()
end

test_24()
```

Output:
```
0.5543628440854502
#<Nyaplot::Frame:0x00000000034c2360 @properties={:panes=>[#<Nyaplot::Plot:0x00000000034f5738 @properties={:diagrams=>[#<Nyaplot::Diagram:0x00000000034d06e0 @properties={:type=>:line, :options=>{:x=>:x, :y=>:y}, :data=>"fc8d361d-9291-4912-b2ba-e258574cff14"}, @xrange=[0, 0], @yrange=[0.0, 1.4000122625885827]>, #<Nyaplot::Diagram:0x00000000034c2b58 @properties={:type=>:line, :options=>{:x=>"data0", :y=>"data1", :color=>:gray}, :data=>"3ceed11a-3684-4d1d-b58f-c8074276a849"}, @xrange=[0, 1], @yrange=[0, 1]>], :options=>{:x_label=>"False Positive Rate", :y_label=>"True Positive Rate", :zoom=>true, :width=>700, :xrange=>[0, 1], :yrange=>[0.0, 1.4000122625885827]}}>], :data=>{"fc8d361d-9291-4912-b2ba-e258574cff14"=>#<Nyaplot::DataFrame:0x00000000034d0d20 @name="fc8d361d-9291-4912-b2ba-e258574cff14", @rows=[{:x=>0, :y=>0.0}, {:x=>0, :y=>0.010309278350515464}, {:x=>0, :y=>1.4000122625885827}]>, "3ceed11a-3684-4d1d-b58f-c8074276a849"=>#<Nyaplot::DataFrame:0x00000000034c3350 @n
```

## Question 3.1 (10 Points)

Implement the linear kernel where $x_i$ and $x_j$ are rows in a dataset.

<!-- grade_id: cell-fcf79802d7f8c72d, points: 0, grade: false, solution: true -->
```ruby
class LinearKernel
  def func x_i, x_j
    # BEGIN YOUR CODE
    return dot(x_i["features"], x_j["features"])
    #END YOUR CODE
  end
end
```

Output:
```
:func
```

<!-- grade_id: cell-3342f2f3fed428da, points: 5, grade: true, solution: false -->
```ruby
def test_31()
  linear = LinearKernel.new
  
  kij = linear.func({"features" => {"a" => 2, "b" => 3}}, {"features" => {"b" => 7, "c" => 1}})  
  assert_in_delta 21, kij, 1e-2, "1"
  
  kij = linear.func({"features" => {"a" => 2, "b" => 3}}, {"features" => {"a" => -5, "b" => 2}})  
  assert_in_delta -4, kij, 1e-2, "1"
end

test_31()
```

## Question 3.2 (10 Points)

Implement the Gaussian Kernel, taking a bandwidth parameter $\sigma$. 

<!-- grade_id: cell-bd59856809fb8bc2, points: 0, grade: false, solution: true -->
```ruby
class GaussianKernel
  def initialize width
    @width = width
  end
  def func x_i, x_j
    # BEGIN YOUR CODE
    diff = Hash.new{|h,k| h[k] = 0.0}
    x_i["features"].each_key do |key|
      if x_j["features"].has_key?(key)
        diff[key] = x_i["features"][key] - x_j["features"][key]
      end
    end
    return Math.exp(-@width * dot(diff, diff))
    #END YOUR CODE
  end
end
```

Output:
```
:func
```

<!-- grade_id: cell-ec6c13625f9e6801, points: 10, grade: true, solution: false -->
```ruby
def test_32()
  spiral = spiral_dataset()
  kernel = GaussianKernel.new 1.0
  
  test_v = kernel.func spiral["data"][0], spiral["data"][2]
  assert_true(test_v > 0.19, "1")
  assert_true(test_v < 0.25, "2")
end

test_32()
```

## Question 3.3 (10 Points): Gram Matrix

Implement the Gram Matrix as a preprocessing step which is first trained and then used to transform a dataset. Given a training set $x_i \in X$ and a kernel function $k$ the Gram Matrix is defined as:

# $G_{i,j} = k(x_i, x_j)$

For any new datapoint z, the _Gram vector_ is defined as follows:

# $g(z) = \left( k(z, x_i), \dots, k(z, x_n) \right)$

Note that on the training $g(x_i) = G_{i,\cdot}$. 

1. Implement a ```train``` method that caches data for evaluation. 
1. Implement the ```func``` method that takes a single row and the gram vector as a row 

<!-- grade_id: cell-6dec8ec0ea6c2f66, points: 0, grade: false, solution: true -->
```ruby
class GramMatrix
  def initialize kernel
    @kernel = kernel
    @data = Array.new
  end
  
  def train data
    # BEGIN YOUR CODE
    @data = data.clone
    @data = data.collect do |r|
      u = r.clone
      u["features"] = r["features"].clone
      u
    end
    return @data
    #END YOUR CODE
  end
  
  def transform z
    # BEGIN YOUR CODE
    gd = Hash.new {|h, k| h[k] = 0.0}
    gd["label"] = z["label"]
    gd["features"] = Hash.new {|h, k| h[k] = 0.0}
    n = @data.size
    n.times do |i|
      x = "k_" + i.to_s
      gd["features"][x] = @kernel.func(@data[i],z)
    end
    return gd
    #END YOUR CODE
  end
end
```

Output:
```
:transform
```

<!-- grade_id: cell-f9dc67d1e3802453, points: 10, grade: true, solution: false -->
```ruby
def test_33()
  lin = LinearKernel.new
  data =  [
    {"features" => {"a" => 2, "b" => 3}, "label" => 1}, 
    {"features" => {"b" => 7, "c" => 1}, "label" => 7},
    {"features" => {"a" => 2, "b" => 3}, "label" => 2}
  ]
  
  z = {"features" => {"a" => -5, "b" => 2}} 

  gmat = GramMatrix.new lin
  gmat.train data

  k_x1 = gmat.transform data[1]
  assert_equal 3, k_x1["features"].size, "1"
  assert_equal 7, k_x1["label"], "2"
  assert_equal dot(data[1]["features"], data[1]["features"]), k_x1["features"]["k_1"], "3"
  assert_equal 21, k_x1["features"]["k_0"], "4"
  k_x2 = gmat.transform z
  assert_equal 3, k_x1["features"].size, "5"
  assert_in_delta -4, k_x2["features"]["k_0"], "6"
end

test_33()
```

## Question 4.1 (10 Points)

Test the Hinge Loss with Linear Kernel on the spiral dataset, return the false positive, true positive, and auc values. Do you see any improvements with this kernelized version? 

<!-- grade_id: cell-6f585f8a01072b23, points: 1, grade: true, solution: true -->
```ruby
def train_linear_kernel_hinge(dataset)
  lr = HingeLossClassifier.new 0.1
  w = Hash.new {|h,k| h[k] = (rand * 0.1) - 0.05}
  
  # BEGIN YOUR CODE
  lk = LinearKernel.new
  gm = GramMatrix.new lk
  gm.train dataset["data"]
  data = dataset["data"].clone
  data = dataset["data"].collect do |x|
    {"features" => (gm.transform x)["features"], "label" => x["label"]}
  end

  sgd = StochasticGradientDescent.new lr, w, 0.1
  50.times do |t|
    sgd.update data.sample(20)
  end
  scores = binary_classification(data, w, lr)
  return scores 
  #END YOUR CODE
end
```

Output:
```
:train_linear_kernel_hinge
```

<!-- grade_id: cell-444aa68056e01128, points: 10, grade: true, solution: false -->
```ruby
def test_41()
  spiral = spiral_dataset()
  scores = train_linear_kernel_hinge(spiral)
  fp, tp, auc = roc_curve(scores)
  puts auc

  assert_true(auc > 0.45, "AUC > 0.45")
  assert_true(auc < 0.7, "AUC < 0.7")    

  plot = Daru::DataFrame.new({x: fp, y: tp}).plot(type: :line, x: :x, y: :y) do |plot, diagram|
    plot.x_label "False Positive Rate"
    plot.y_label "True Positive Rate"
  end
  plot.add(:line, [0,1], [0,1]).color(:gray)
  plot.show()
end

test_41()
```

Output:
```
0.555106812626209
#<Nyaplot::Frame:0x0000000003fafc50 @properties={:panes=>[#<Nyaplot::Plot:0x0000000003fb5b28 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000000003fb1be0 @properties={:type=>:line, :options=>{:x=>:x, :y=>:y}, :data=>"fc7d00c9-12f5-4d89-b7cc-a614eb1e0c04"}, @xrange=[0, 0], @yrange=[0.0, 91.18039060582247]>, #<Nyaplot::Diagram:0x0000000003fb0510 @properties={:type=>:line, :options=>{:x=>"data0", :y=>"data1", :color=>:gray}, :data=>"1cbb9110-e32a-477e-b102-cf487630022e"}, @xrange=[0, 1], @yrange=[0, 1]>], :options=>{:x_label=>"False Positive Rate", :y_label=>"True Positive Rate", :zoom=>true, :width=>700, :xrange=>[0, 1], :yrange=>[0.0, 91.18039060582247]}}>], :data=>{"fc7d00c9-12f5-4d89-b7cc-a614eb1e0c04"=>#<Nyaplot::DataFrame:0x0000000003fb2518 @name="fc7d00c9-12f5-4d89-b7cc-a614eb1e0c04", @rows=[{:x=>0, :y=>0.0}, {:x=>0, :y=>0.010309278350515464}, {:x=>0, :y=>91.18039060582247}]>, "1cbb9110-e32a-477e-b102-cf487630022e"=>#<Nyaplot::DataFrame:0x0000000003fb0d58 @name=
```

## Question 3.3 (5 points)

Is the AUC for Linear Kernel much better than before? Why or Why not?

<!-- grade_id: cell-6cf57bec611f2fa2, points: 0, grade: false, solution: true -->
```ruby
def train_gaussian_kernel_hinge(dataset)
  lr = HingeLossClassifier.new 0.1
  w = Hash.new {|h,k| h[k] = (rand * 0.1) - 0.05}
  
  # BEGIN YOUR CODE
  kernel = GaussianKernel.new 1.3
  gm = GramMatrix.new kernel
  gm.train dataset["data"]
  data = dataset["data"].clone
  data = dataset["data"].collect do |x|
    {"features" => (gm.transform x)["features"], "label" => x["label"]}
  end

  sgd = StochasticGradientDescent.new lr, w, 1.0
  50.times do |t|
    sgd.update data.sample(20)
  end
  scores = binary_classification(data, w, lr)
  return scores 
  #END YOUR CODE
end
```

Output:
```
:train_gaussian_kernel_hinge
```

<!-- grade_id: cell-e94797f3c52a8594, points: 10, grade: true, solution: false -->
```ruby
def test_42()
  spiral = spiral_dataset()
  scores = train_gaussian_kernel_hinge(spiral)
  fp, tp, auc = roc_curve(scores)
  puts auc

  assert_true(auc > 0.5, "AUC > 0.5")
  assert_true(auc > 0.6, "AUC > 0.6")  
  assert_true(auc > 0.7, "AUC > 0.7")    
  assert_true(auc > 0.8, "AUC > 0.8")    
  assert_true(auc < 1.0, "AUC < 1.0")    

  plot = Daru::DataFrame.new({x: fp, y: tp}).plot(type: :line, x: :x, y: :y) do |plot, diagram|
    plot.x_label "False Positive Rate"
    plot.y_label "True Positive Rate"
  end
  plot.add(:line, [0,1], [0,1]).color(:gray)
  plot.show()
end

test_42()
```

Output:
```
0.8872356254649806
#<Nyaplot::Frame:0x0000000003ab8580 @properties={:panes=>[#<Nyaplot::Plot:0x0000000003acaaa0 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000000003aba920 @properties={:type=>:line, :options=>{:x=>:x, :y=>:y}, :data=>"2c09eb7d-f1c9-4536-848e-5d80299705a5"}, @xrange=[0, 0], @yrange=[0.0, 1.0334493851344864]>, #<Nyaplot::Diagram:0x0000000003ab8d28 @properties={:type=>:line, :options=>{:x=>"data0", :y=>"data1", :color=>:gray}, :data=>"aaa1ec8d-6f5b-4d5c-9a55-d36d6ec07290"}, @xrange=[0, 1], @yrange=[0, 1]>], :options=>{:x_label=>"False Positive Rate", :y_label=>"True Positive Rate", :zoom=>true, :width=>700, :xrange=>[0, 1], :yrange=>[0.0, 1.0334493851344864]}}>], :data=>{"2c09eb7d-f1c9-4536-848e-5d80299705a5"=>#<Nyaplot::DataFrame:0x0000000003abb3c0 @name="2c09eb7d-f1c9-4536-848e-5d80299705a5", @rows=[{:x=>0, :y=>0.0}, {:x=>0, :y=>0.010309278350515464}, {:x=>0, :y=>1.0334493851344864}]>, "aaa1ec8d-6f5b-4d5c-9a55-d36d6ec07290"=>#<Nyaplot::DataFrame:0x0000000003ab9598 @n
```

## Question 5.1 (7 Points)

Compare the un-kernelized version, linear kernel, and Gaussian kernels on this dataset. Provide answers to the following questions as an array below:

1. Which model had the best performance as measured by AUC
 * (A) Linear Kernel
 * (B) Gaussian Kernel
 * (C) Non-Kernel
 * (D) They were the same
1. If there are $N$ examples in the training set, what is the size of the Gram matrix during training?
 * (A) $O(N)$
 * (B) $O(N \log N)$
 * (C) $O(1)$
 * (D) $O(N^2)$
1. For linear classifier, if there are $N$ examples with $d$ features in the training set, what is the time complexity of scoring one example?
 * (A) $O(d)$
 * (B) $O(N^2)$
 * (C) $O(Nd)$
 * (D) $O(N \log d)$
1. For kernelized classifier, if there are $N$ examples with $d$ features in the training set, what is the time complexity of scoring one example?
 * (A) $O(N \log d)$
 * (B) $O(d)$
 * (C) $O(N^2)$
 * (D) $O(Nd)$
1. Why do you think the Gaussian kernel would have better performance for the spiral dataset?
 * (A) Hinge loss causes the linear kernel to overfit, but the Gaussian kernel is very smooth and therefore less prone to overfitting.
 * (B) Gaussian kernel did not perform better than anything else.
 * (C) The spiral dataset is non-linear and the Gaussian kernel can learn non-linear boundaries.
 * (D) The data follows a Gaussian distribution therefore Gaussian kernels should fit the best. 
 
 
If your answers were all A, A, A, A, A, then write the following:

```ruby
def answer_51()
  answers = ["A", "A", "A", "A", "A"]
  return answers
end
```

<!-- grade_id: cell-0e8a77df1540314c, points: 0, grade: false, solution: true -->
```ruby
def answer_51()
  # BEGIN YOUR CODE
  answers = ["B", "D", "C", "D", "C"]
  #END YOUR CODE
  return answers
end
```

Output:
```
:answer_51
```

<!-- grade_id: cell-f64f4eb643d0d952, points: 7, grade: true, solution: false -->
```ruby
t51_answers = answer_51()

assert_not_nil t51_answers, "1"
assert_true(t51_answers.is_a?(Array))
assert_equal(5, t51_answers.size)
assert_true(t51_answers.any? {|a| a.size == 1 and a =~ /[A-Z]/})

```

```ruby

```